Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below.

Rename this problem sheet as follows:

    ps{number of lab}_{your user name}_problem{number of problem sheet in this lab}
    
for example
    
    ps2_blja_problem1

Submit your homework within one week until next Monday, 9 a.m.

```python
NAME = ""
EMAIL = ""
USERNAME = ""
```

---

# Introduction to Data Science
## Lab 3: Regression

The previous exercises gave an introduction to Python, Numpy and Pandas. Beginning with this exercise, we shift our focus to statistical learning itself. To this end, we will employ the module scikit-learn which offers many functions we will cover over the remaining semester.

### Part A: Linear regression for the Advertising dataset

If not already done, please download the file [Advertising.csv](https://www.tu-chemnitz.de/mathematik/numa/lehre/ds-2019/Exercises/ps3/Advertising.csv) and move it into the current directory on the Jupyter Hub.

**Task**: Import the *Advertising* dataset using `pandas` `read_csv` function. Ensure, that the first column is treated as the index column.

<!-- grade_id: cell-8d64ab6b1a68631f, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

# YOUR CODE HERE
raise NotImplementedError()

# Print first entries of adv
print(adv.head(3))
```

<!-- grade_id: cell-b1115ea7f17d9f4b, points: 1, grade: true, solution: false -->
```python
assert len(adv.index) == 200
assert adv.TV.mean() == 147.0425
assert adv.shape == (200,4)
```

For convenience, we extract the values from this `pandas`-DataFrame.
This can be done with the method `values`, which returns a `numpy` array.

**Task**: Extract the data in the following manner:
- `X` contains all predictor variables, i.e., the values from the columns `TV`, `Radio` and `Newspaper`.
- `Y` contains the dependent variable `Sales`

<!-- grade_id: cell-f0d1c2f7e6986c87, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6af75eca942f8185, points: 1, grade: true, solution: false -->
```python
assert X.shape == (200,3)
assert Y.shape == (200,)
assert type(X) == np.ndarray
assert type(Y) == np.ndarray
```

Note that the variables `X` and `Y` contain only the values itself, and no further information like the column title.

Using the `numpy` function `hsplit`, we can split an array horizontally.
This will become handy in many circumstances.
Here, we split the 2-dimensional numpy array `X` into three 
**2**-dimensional slices `tv`, `radio` and `news`.
Note that the returned arrays have still the second dimension (which is one)!

```python
tv, radio, news = np.hsplit(X,3)
```

Compute for each of the 3 predictor variables **TV**, **radio** and **newspaper** simple (1-dimensional) linear regressions, e.g.

$$ y^{TV}_i \approx \beta_0^{TV} + \beta_1^{TV} \, x_i^{TV}$$

Use the following function:

    from sklearn.linear_model import LinearRegression
    
Store the intercepts, i.e., the values $\beta_0$ in variables
       
    intercept_tv, intercept_radio and intercept_news
    
and the linear coefficients, i.e., the values $\beta_1$ in variables

    lincoef_tv, lincoef_radio and lincoeff_news

To print your results in a nice fashion, you can use a command similar to

    print('y = %5.4f + %5.4f x TV' % (intercept, lincoef))

<!-- grade_id: cell-5cb14fe0e222492c, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ed24b27d868aff40, points: 1, grade: true, solution: false -->
```python
assert np.abs(intercept_tv-7.0326) < 1e-4
assert np.abs(lincoef_tv-0.0475) < 1e-4
```

You should observe, that the regression coefficients for **TV** and **newspaper** are very similar.
As you already know from the lecture, it is not satisfying from a mathematical point of view to restrict our investigation to the absolute values of the coefficients.


### Part B: Assessing the quality of a linear fit.

In the lecture you learned about different measures for assessing the quality of a linear fit.
In the last exercise, we already implemented a function to compute the mean squared error (MSE).

This time, we want to compare the $R^2$ scores. You can use the method `score()` of a `LinearRegression` to get the $R^2$ values.
Remember that this value is the proportion of variability in $Y$ explained using **TV**, **radio** or **newspaper** as predictor in a 1-dimensional linear regression fit.

**Task**: Compute the $R^2$ scores and store them in variables
    
    R2_tv, R2_radio, R2_news

<!-- grade_id: cell-fd3b6790a9e543c0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print("R^2 for TV:        ", R2_tv)
print("R^2 for radio:     ", R2_radio)
print("R^2 for newspaper: ", R2_news)
```

<!-- grade_id: cell-b9cdf81fe534585a, points: 1, grade: true, solution: false -->
```python
assert np.abs(R2_tv - 0.611875050850071) < 1e-9
```

### Part C: Predicting values

Now we want to compute the predicted value of sales if we restrict our prediction to one input, i.e. **TV**, **radio** or **newspaper**, resp.
Predict the values $\hat{y}^{TV}$ $\hat{y}^{radio}$ and $\hat{y}^{newspaper}$ using the `LinearRegression`-method `predict()` and store them in variables

    y_tv, y_radio, y_news

<!-- grade_id: cell-9e5459d7aaa2dfc8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-71f652a968725ea2, points: 1, grade: true, solution: false -->
```python
assert np.abs(y_tv.mean() -14.0225) < 1e-5
assert np.abs(y_tv.std() - 4.071006120646744) < 1e-5
```

### Part D: Plotting
Plot the datapoints as well as the corresponding regression line for each of the inputs  **TV**, **radio** or **newspaper**.

You can use the functions `subplots` or `fig.add_subplot` to arrange the plots in one figure.

<!-- grade_id: cell-f5323561f0d02b4a, points: 1, grade: true, solution: true -->
```python
# We plot our findings using subplots
import matplotlib.pyplot as plt
plt.rcParams['figure.figsize'] = (16,9)

fig = plt.figure()
fig.add_subplot(1,3,1)

# YOUR CODE HERE
raise NotImplementedError()

plt.title('Lin. regr. TV')

fig.add_subplot(1,3,2)

# YOUR CODE HERE
raise NotImplementedError()

plt.title('Lin. regr. radio')

fig.add_subplot(1,3,3)

# YOUR CODE HERE
raise NotImplementedError()

plt.title('Lin. regr. newspaper');

```

### Part E: Statistical functions

Take a closer look at the correlation matrix for the `DataFrame` `adv`.
You can use the method `corr()` that is implemented for pandas 
`DataFrames`.
Which features are correlated most strongly?

```python
adv.corr()

```

Name the two features, that are correlated most strongly in the next cell!

<!-- grade_id: cell-3d9e79b4a9566134, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Investigate the statistical significance of the medium **newspaper** in a linear regression involving only this feature. Use a **t-test** for this purpose as described on slide 16 in the document 03_IntroDS_Regression.pdf of the lecture notes.

You should observe the following values:

|Coefficient | Estimate | SE | t-statistic | p-value|
|:-----------|----------|----|-------------|--------|
| $\beta_0$  | 12.351 | 0.621 | 19.88 | < 0.0001 |
| $\beta_{newspaper}$ | 0.055 | 0.017 | 3.30 | 0.00115


You should use `scipy` to get the $t$-distribution using

    from scipy.stats import t
    
The cumulative distribution function at a point `x` for `n` degrees of freedom can than be called by

    t.cdf(x, n)

<!-- grade_id: cell-02461ff24a250328, points: 3, grade: true, solution: true -->
```python
from scipy.stats import t

# YOUR CODE HERE
raise NotImplementedError()

# Print information
print("Intercept: %7.4f" % beta_0)
print("Lin. coef: %7.4f" % beta_1)
print("")
print("std-error Intercept: %7.4f" %  SE0)
print("std-error Lin. coef: %7.4f" % SE1)
print("")
print("t-statistic Intercept: %6.2f" % tstat_0)
print("t-statistic Lin. coef: %6.2f" % tstat_1)
print("")
print("P-value Intercept: %7.5f" % pval_0)
print("P-value Lin. coef: %7.5f" % pval_1)
```

### Part F: Linear regression on all prectictors

Now construct a linear regression on all three predictor variables, i.e.

$$y_i ≈ \beta_0 + \beta_{TV} x^{TV}_i + \beta_{radio} x^{radio}_i + \beta_{newspaper} x^{newspaper}_i$$

Ensure that the intercept and linear regression coefficients are stored in the variables

    beta_0, beta_tv, beta_radio and beta_news

<!-- grade_id: cell-59ddc25633526b2d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
print('y = %5.4f + %5.4f x TV + %5.4f x radio + %5.4f x newspaper' % (beta_0, beta_tv, beta_radio, beta_news))
```

<!-- grade_id: cell-cd23d2ff6972429b, points: 1, grade: true, solution: false -->
```python
assert np.abs(beta_0 - 2.9389) < 1e-4
```

What portion of the variance (between 0 and 1) is explained by this linear regression fit? Store your answer in the variable `explained_var`.

<!-- grade_id: cell-9cb2dad97c8c9bf5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-66b0fe342bdf5b49, points: 1, grade: true, solution: false -->
```python
print('The portion of the variance that can be explained by the full model is about %8.6f' % explained_var)

```

Now perform a linear regression that incorporates only the predictors **TV** and **radio**.
Compute the $R^2$-value of this linear regression model and compare it to the $R^2$-value of the full multiple linear regression.

<!-- grade_id: cell-caee2298f250d394, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-664727b9fc60e7e8, points: 1, grade: true, solution: false -->
```python
print('The R^2 score for the model using TV and Radio as predictors is %8.6f' % R2_tv_radio)

```

You should observe that the $R^2$-score for the linear regression fit incorporating all three features is only marginally larger than the score using only **TV** and **radio** for the prediction.
Thus, it might be sufficient to exclude the **feature** newspaper from our prediction.

The procedure we did today is called *feature selection* is should be one of the first steps in every statistical learning problem.

### Part G: Computing the t-statistic for the full model

We have already seen, that the **t-test** comes in handy when one has to decide whether a coefficient for a single feature is significant or not.
As has been outlined in the lecture, one can also use the t-test in a multiple linear regression fit

$$ Y = X \beta + \varepsilon $$

while the intercept is incorporated into $X$, i.e. a column containing only ones is stacked in front of the original matrix $X$.

The formula to compute the test statistic in this generalized setting is

$$ t_j = \frac{\hat{\beta}_j}{\hat{\sigma} \sqrt{v_j}} $$

while $\hat \beta_j$ is the $j$-th entry of the coefficient vector

$$ \hat \beta = (X^\top X)^{-1} X^\top y, $$

$\hat{\sigma}$ is the unbiased estimate of $\sigma$, which is determined by

$$ \hat{\sigma} = \sqrt{\frac{1}{n-p-1} \, \sum_{i=1}^n (y_i - \hat{y}_i)^2} $$

and $v_j$ is the $j$-th diagonal element of the matrix $(X^\top X)^{-1}$.

Then $t_j$ is distributed according to a $t$-distribution with $n-p-1$ degrees of freedom (dofs). 

**Task**: Compute the values in the following statistic and try to print it in a similar way. 

|  Coefficient     | Estimate  | SE    | t-statistic | p-value |
|:-----------------|-----------|-------|-------------|---------|
| $\beta_0$        | 2.939     |0.3119 | 9.42        | < 0.0001|
| $\beta_{TV}$     | 0.046     |0.0014 | 32.81       | < 0.0001|
| $\beta_{radio}$  | 0.189     |0.0086 | 21.89       | < 0.0001|
| $\beta_{news}$   | −0.001    |0.0059 | −0.18       | 0.8599  |




<!-- grade_id: cell-5de872312430d808, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```