Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below.

Rename this problem sheet as follows:

    ps{number of lab}_{your user name}_problem{number of problem sheet in this lab}
    
for example
    
    ps2_blja_problem1

Submit your homework within one week until next Monday, 9 a.m.

```python
NAME = ""
EMAIL = ""
USERNAME = ""
```

---

# Introduction to Data Science
## Lab 2: Data import and linear regression

This course assumes that you are comfortable with the basic functions of Jupyter and Python.

### Part A: Introduction Data Import

Download the file `Advertising.csv` from the homepage’s exercise section and upload it to your Jupyter Hub folder.
Take a short look at the `csv`-file using a spreadsheet, e.g., LibreOffice.
The file contains information about the sales of products in different markets, along with advertising budgets in the three media: **TV**, **radio** and **newspaper**.

**Task**: Import the `csv`-file using the `numpy` function `genfromtxt` and store it as an array `X`.

<!-- grade_id: cell-824784e084c340f0, points: 0, grade: false, solution: true -->
```python
import numpy as np

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-784ee0c212e8cdfc, points: 1, grade: true, solution: false -->
```python
assert X.shape == (200,5)
assert X[34][1] == 95.7
```

**Task**: Extract the columns from the array `X` and store them as 1-dimensional arrays `idx`, `tv`, `radio`, `newspaper` and `sales`, e.g.,
    
    idx = X[:, 0]

<!-- grade_id: cell-983eb3ebbf4d3eed, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-05f112d846deb52b, points: 1, grade: true, solution: false -->
```python
assert idx.ndim == 1
assert tv.ndim == 1
assert radio.ndim == 1
assert newspaper.ndim == 1
assert sales.ndim == 1
i=27
assert idx[i] == 28
assert tv[i] == 240.1
assert radio[i] == 16.7
assert newspaper[i] == 22.9
assert sales[i] == 15.9
```

**Task**: Add subplots to plot sales against radio as well as sales against newspaper.

<!-- grade_id: cell-cc1107edacb5af03, points: 1, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt
plt.rcParams['figure.figsize'] = (16, 9)
fig1 = plt.figure()
fig1.add_subplot(1,3,1)
plt.plot(tv, sales, 'ro')
plt.xlabel('TV budget')
plt.ylabel('sales')
plt.title('TV ads')

fig1.add_subplot(1,3,2)
# TASK: Plot sales against radio
# YOUR CODE HERE
raise NotImplementedError()

fig1.add_subplot(1,3,3)
# TASK: Plot sales against newspaper
# YOUR CODE HERE
raise NotImplementedError()

plt.show()
```

### Part B: Creation of a function in Python
The mean squared errer, short **MSE**, is one of the most important performance indicators for the quality of data fitting.
The goal of this exercise is to implement the function `computeMSE` with the following **input**:
- the observations $y_i \in Y$, $i = 1, \ldots, N$ that belong to measurements $x_i \in X$, $i = 1, \ldots, N$
- the predictions of $f(x_i)$, which are denoted by $\hat f(x_i)$, $i = 1, \ldots, N$

and corresponding **output**:

$$
    MSE = \frac{1}{N} \sum_{i=1}^N (y_i - \hat f (x_i))^2
$$

<!-- grade_id: cell-656779d5a9b31dc3, points: 0, grade: false, solution: true -->
```python
import numpy as np

# Define function for mean squared error
def computeMSE(y, fhatx):
    # YOUR CODE HERE
    raise NotImplementedError()
    return mse
```

<!-- grade_id: cell-103b708d111d04d3, points: 2, grade: true, solution: false -->
```python
x = np.array([0.55, 0.72, 0.6 , 0.54, 0.42,
    0.65, 0.44, 0.89, 0.96, 0.38,
    0.79, 0.53, 0.57, 0.93, 0.07,
    0.09, 0.02, 0.83, 0.78, 0.87])

y = np.array([6.9 , 7.58, 7.03, 6.61, 5.84,
    7.32, 6.29, 8.38, 9.03, 5.75,
    7.95, 6.63, 7.  , 8.8 , 4.37,
    4.49, 4.01, 7.95, 7.87, 8.37])

fhatx = np.array([
       6.74792024, 7.61454115, 7.00280875, 6.69694254, 6.08521014,
       7.25769725, 6.18716554, 8.48116205, 8.83800595, 5.88129934,
       7.97138505, 6.64596484, 6.84987564, 8.68507285, 4.30099064,
       4.40294604, 4.04610214, 8.17529585, 7.92040735, 8.37920665])

print('MSE: ', computeMSE(y, fhatx))

assert np.abs(computeMSE(y, fhatx) - 0.01375124894963537) < 1e-9
```

**Task**: Compute the linear regression coefficients $\beta_0, \beta_1 \in \mathbb{R}$ with the `numpy` function `np.polyfit`.
Remember, linear regression finds the values $\beta_0, \beta_1 \in \mathbb{R}$ which solve the minimization problem

$$
\text{Minimize } \frac{1}{N} \sum_{i=1}^N \left( y_i - (\beta_0 + \beta_1 x_i) \right)^2 \text{ over } \beta_0, \beta_1 \in \mathbb{R}
$$

*Hint*: Again, use the question mark `?` to get help and have a look at the documentation.

<!-- grade_id: cell-4f6c481b7e4d76bf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-515608379fe4789e, points: 1, grade: true, solution: false -->
```python
assert np.abs(beta[0] - 5.09777001) < 1e-8
assert np.abs(beta[1] - 3.94414674) < 1e-8
```

**Task**: Compute the prediction of `x` by yourself and store it in a variable `z`.
The values should coincide with those of the variable `fhatx` (by at most `1e-8`). You can use the function `np.polyval`.

<!-- grade_id: cell-d190b7c4862b9971, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c6db1531a735649e, points: 1, grade: true, solution: false -->
```python
assert np.max(np.abs(fhatx - z)) < 1e-8
```

**Task**: Plot the data points $(x_i, y_i)$, $i = 1 \ldots, N$ together with the least squares line aka regression line.

<!-- grade_id: cell-ae78dcf4fb4cbe76, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Part C: Introduction Pandas
In **Part A**, you got to know a method to import `csv`-files using the function `np.genfromtxt`.
At one point or another, we would have to deal with a problem
that is inherent to numpy-arrays, namely that numpy-arrays can only handle one data type at a time.
If we have different kinds of data like booleans, floats, integers or
strings, we have to take a different route.
One possible solution lies in the usage of the package `csv`.
Here, every single row is scanned seperately, and thus can be handled to catch special cases.
Another possibility is to use the package `pandas`, whose complexity is between the other two. It can be imported by

    import pandas as pd

and `csv`-files can be imported by the function `pd.read_csv`.

**Task**: Work through the [pandas tutorial](https://pandas.pydata.org/pandas-docs/stable/10min.html#min) (this is a link!).

**Task**: Download the file `Auto.csv` from the lecture’s homepage. Import the `csv`-file using the pandas function `read_csv` as a `DataFrame` named `Auto`.
Beware of the missing values in the `csv`-file.
You can use the optional parameter `na_values` from the function `read_csv`.
In this problem, we want to **remove those data sets** that contain missing values.
You should use the method `dropna(axis=0, inplace=True)` for this purpose.

<!-- grade_id: cell-efd646077862b192, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3f85fae5600a9d5b, points: 1, grade: true, solution: false -->
```python
assert Auto.shape == (392,9)
assert Auto.iloc[301,2] == 85
assert np.abs(Auto["mpg"].mean() - 23.445918367346938) < 1e-8
```

**Task**: Create a short summary of the most important statistics of the data set using the method `describe`.

<!-- grade_id: cell-73e33597f66d3327, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Task**: Create a graphical overview of the distributions of the input variables of the data set using the method `hist`.

<!-- grade_id: cell-2ccdaf2c47ba29e9, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Task**: Create a scatter matrix using the function `pd.plotting.scatter_matrix` for the variables `["horsepower", "mpg", "weight"]`.

<!-- grade_id: cell-9c3cea01f6163f40, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Create two figures that relate the variable `"horsepower"` with `"mpg"` and `"weight"`, resp.
Use the possibilities that are provided by pandas.

The following code cell investigates a linear or quadratic connection between horsepower and mpg (miles per gallon).

```python
# Investigate linear and quadratic connection between horsepower and mpg
x = Auto.horsepower
y = Auto.mpg

mpgbeta1 = np.polyfit(x,y,deg=1)
mpgbeta2 = np.polyfit(x,y,deg=2)

Auto.plot(x='horsepower', y = 'mpg', marker='o', alpha = .7, kind='scatter')
xr = np.linspace(Auto.horsepower.min(), Auto.horsepower.max(), 100)
plt.plot(xr,np.polyval(mpgbeta1,xr), c='r')
plt.plot(xr,np.polyval(mpgbeta2,xr), c='b')

print('Investigating mpg against horsepower')
print('MSE for linear fit: ', computeMSE(y, np.polyval(mpgbeta1, x)))
print('MSE for quadratic fit: ', computeMSE(y, np.polyval(mpgbeta2, x)))

```

**Task**: Use the example from above and investigate a possible linear and quadratic relationship between horsepower and weight.

<!-- grade_id: cell-b75fc32748a1111f, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

```python

```