Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below.

Rename this problem sheet as follows:

    ps{number of lab}_{your user name}_problem{number of problem sheet in this lab}
    
for example
    
    ps8_blja_problem1

Submit your homework until Thursday, December 17, 2020, 9 am.

```python
NAME = ""
EMAIL = ""
USERNAME = ""
```

---

# Introduction to Data Science
## Lab 8: Linear discriminant analysis (LDA) and Logistic Regression for the classification of iris types

The flower petal data set consists of 3 different types of irises’ (Setosa, Versicolour, and Virginica) petal and sepal size.

Given the predictors
- 1st column: sepal length,
- 2nd column: sepal width,
- 3rd column: petal length and
- 4th column: petal width,

our goal is to predict the correct class (0-Setosa, 1-Versicolour or 2-Virginica).

The data set is part of `scikit-learn`'s datasets module and can be imported with the following commands:

```python
from sklearn import datasets
iris = datasets.load_iris()
```

The data comes as a dictionary. You can access the predictors using `iris.data` and the classes using `iris.target`.

**Task (1 point)**: Store the predictors in a variable `X` and the response in a variable `y`.

<!-- grade_id: cell-4c5d0ba46cdd1586, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-53d8fde18f295bd5, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150,4)
assert y.shape == (150,)
```

**Task (2 points)**: Plot the sepal length on the x-axis and the sepal width on the y-axis. Color each of the three types of irises differently.
Add a legend that gives the correct iris type (0-Setosa, 1-Versicolour, 2-Virginica)

<!-- grade_id: cell-a99e46c9b162136b, points: 2, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt
%matplotlib inline
# YOUR CODE HERE
raise NotImplementedError()
```

**Task (2 points)**:
Split your data into a training and a test set.
Put the first 40 samples within each class in the training set and the remaining samples in a test data set.

Store the training set in variables `Xtrain` and `ytrain`, and the test set in variables `Xtest` and `ytest`, resp.

<!-- grade_id: cell-3397ff5382060b53, points: 0, grade: false, solution: true -->
```python
import numpy as np
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4cff00529ce30e6d, points: 2, grade: true, solution: false -->
```python
assert Xtrain.shape == (120,4)
assert Xtest.shape == (30,4)
assert ytrain.shape == (120,)
assert ytest.shape == (30,)
assert abs(Xtrain.mean() - 3.485208333333333) < 1e-10
assert abs(ytrain.mean() - 1) < 1e-10
assert abs(ytest.mean() - 1) < 1e-10
assert abs(Xtest.mean() - 3.3816666666666673) < 1e-10
```

### Part B: Linear discriminant analysis

In the lecture you've heard about the classification method called
*Linear discriminant analysis (LDA)*.

**Task (1 points)**: Find a way using `scikit-learn` to accomplish a linear discriminant analysis on the **training data set**.

Perform an LDA using only the first two predictors, i.e., `sepal length` and `sepal width`.
Store your trained model in the variable `lda`.

<!-- grade_id: cell-02fee2efc178e8e6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f84d0268df55c5ca, points: 1, grade: true, solution: false -->
```python
assert 'lda' in locals()
assert abs(lda.predict_proba([[2.1,1.1]])[0][0] - 0.7867422283434491) < 1e-10
```

**Task (1 point)**: What is the proportion of correctly classified irises in the **test data set**? Store your answer in the variable `prop1`.

<!-- grade_id: cell-902ceff2ae133619, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1cc1e8c94bb792dd, points: 1, grade: true, solution: false -->
```python
assert 'prop1' in locals()
```

**Task (2 points)**: Now, incorporate all of the predictors and perform a second linear discriminant analysis using **only the training data**.
How does the proportion of correct classifications change (for the **test data**)?
Store the proportion of correct classifications for the test set in the variable `prop2`.

<!-- grade_id: cell-8b7d19220834da29, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e8a557d63a9a00ad, points: 2, grade: true, solution: false -->
```python
assert 'prop2' in locals()
```

### Part C: Multi-class logistic regression

Here, we want to apply logistic regression in a multi-class setting.

#### One vs. rest approach
One way to extend the logistic regression to a setting with $k$ classes is by training not one but $k$ models for $k > 2$, one for each individual class.
As the name suggests, we train one model for each individual class $i$ and try to fit a logistic regression model to a modified data set, where the responses of members belonging to class $i$ are set to `True` and **all** others are set to `False`, i.e. we keep class $i$, set their responses to `True` and modify the responses of the remaining data and set those to `False`.

Fortunately, this **one vs. rest** approach is implemented for many models, and we can train it using the functions which have been already used for the *simple logistic regression* problem.

**Task (1 point)**: Train a logistic regression model (on our **training data**) with the following parameters: 
- penalty parameter: `C = 1e10` 
- solver: `solver = 'liblinear'`
- multi-class option active: `multi_class`

Store your model in the variable `lr`.

<!-- grade_id: cell-719cc6f0de88caea, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LogisticRegression
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0293e9d4d1d99ce4, points: 1, grade: true, solution: false -->
```python
assert 'lr' in locals()
assert abs(lr.predict_proba([[2.1,1.1,1.1,1.0]])[0][0] - 0.07823538026445785) < 1e-6
```

**Task (1 point)**: Store the proportion of correct classifications for the **test set** in the variable `prop0`.

<!-- grade_id: cell-a0082b233f668687, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d632bb14fc6d2528, points: 1, grade: true, solution: false -->
```python
assert 'prop0' in locals()
```