Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below.

Rename this problem sheet as follows:

    ps{number of lab}_{your user name}_problem{number of problem sheet in this lab}
    
for example
    
    ps7_blja_problem1

Submit your homework within one week until next Thursday, December 3, 2020, 9 am.

```python
NAME = ""
EMAIL = ""
USERNAME = ""
```

---

# Introduction to Data Science
## Lab 7: Logistic regression
### Part A: Implementing the logistic function

**Task**:
Implement the logistic function

$$ \sigma(x) = \frac{e^x}{1+e^x} = \frac{1}{1+e^{-x}}$$

using methods from `numpy`.

<!-- grade_id: cell-e34e975989fd0590, points: 0, grade: false, solution: true -->
```python
import numpy as np
def sigma(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9d5048808a942daa, points: 1, grade: true, solution: false -->
```python
assert sigma(0) == 1/2
assert all(sigma(np.ones((4,))) == 1/(1+np.exp(-1)))
```

Next, we want to investigate how the shape of the logistic function changes for an affine linear input, i.e.,

$$ \sigma(\beta_0 + \beta_1 x) $$

for different values of $\beta_0$ and $\beta_1$.

For this purpose, the following function `plot_sigma` (which uses your function `sigma`) has been defined.

**Task**: Understand and execute the following code cell! There will be no plot here, since we are only defining a function.

```python
import matplotlib.pyplot as plt
%matplotlib inline

def plot_sigma(b0, b1):
    plt.figure('Sigma interactive')
    x = np.linspace(-10,10,1001)
    plt.plot(x, sigma(b0 + b1*x), label = 'Logistic function')
    plt.plot(x, 0.5*np.ones(x.shape), label = 'Const 0.5')
    plt.ylim(-0.1, 1.1)
    plt.xlabel('x')
    plt.ylabel('p(x)')
    plt.legend()
    plt.show()
```

**Task**: Call the function `plot_sigma` with arguments `b0 = 0.3` and `b1 = 0.4`.

<!-- grade_id: cell-49a4ec4e6f79756e, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Task**: Execute the next code cell and take your time to investigate how the values of $\beta_0$ and $\beta_1$ affect the behaviour of the shape of the function.

**You have nothing to implement here, only evaluate the cell below and play with the sliders.**

```python
from ipywidgets import interactive
interactive_plot = interactive(plot_sigma, b0=(-10.0, 10.0, 1.0), b1=(-3., 3., 0.2))
output = interactive_plot.children[-1]
output.layout.height = '350px'
interactive_plot
```

### Decide whether the statements are true or false!

Set the answer-variable either to `True` or `False`.

**Statement A**: The logistic function is monotonically increasing for $\beta_0 > 0$.
Store your answer in the variable `answer_A`.

<!-- grade_id: cell-ec99f53493c4e27e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2c3481b9216171ec, points: 1, grade: true, solution: false -->
```python
assert 'answer_A' in locals()
```

**Statement B**: The logistic function is monotonically increasing for $\beta_1 > 0$.
Store your answer in the variable `answer_B`.

<!-- grade_id: cell-cbe3bf01fbd82764, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-38098558fe896986, points: 1, grade: true, solution: false -->
```python
assert 'answer_B' in locals()
```

**Task**: Determine the derivative of $\sigma$ with respect to $x$, i.e., determine $\sigma'(x)$ and write it in the following `Markdown` cell.
Furthermore, confirm that the equation

$$
\sigma'(x) = \sigma(x) \, \left( 1- \sigma(x) \right)
$$

holds.

**Hint**: Use the chain rule!

<!-- grade_id: cell-fd5a2bd92861c285, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE