Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below.

Rename this problem sheet as follows:

    ps{number of lab}_{your user name}_problem{number of problem sheet in this lab}
    
for example
    
    ps7_blja_problem2

Submit your homework within one week until next Thursday, December 3, 2020, 9 am.

```python
NAME = ""
EMAIL = ""
USERNAME = ""
```

---

# Introduction to Data Science
## Lab 7: Logistic regression
### Part B: Logistic regression in practice

In this lab, we want to investigate the `Default` data set known from the lecture.
It contains the predictors
- `student` status, either `'Yes'` or `'No'`
- `balance`, i.e., monthly credit card balance
- yearly `income`
and the response
- `default`, which is either `'Yes'` or `'No'`

We first load the necessary modules.

By the way, the command
    
    plt.rcParams['figure.figsize'] = [13, 5]
    
changes the default size of figures (in inches).

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
plt.rcParams['figure.figsize'] = [13, 5]
```

**Task**: Download the file `Default.csv`.
Read it using the `pandas` function `read_csv` and store the `pandas DataFrame` in the variable `D`.
Make sure that:
- the index column is recognized appropriately.
- the column titles are correct

<!-- grade_id: cell-b1d3abed11dc76fb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d8859f83966b0cca, points: 1, grade: true, solution: false -->
```python
assert 'D' in locals()
assert isinstance(D.index, pd.Int64Index)
assert isinstance(D.columns, pd.Index)
assert D.columns[1] == 'student'
assert D.shape == (10000, 4)
```

**Task**: Inspect the data using the methods you've learned so far, e.g., `describe`, `hist`, `head`, etc.

<!-- grade_id: cell-a3ce086647901439, points: 1, grade: true, solution: true -->
```python
# Task: Apply here at least two different methods to inspect the data set D

# YOUR CODE HERE
raise NotImplementedError()
```

You should observe that the method `describe` only contains the predictors `balance` and `income`, but not  `default` and `student`.

This is due to the fact that these values were read in by the `read_csv` function as **strings**.
We know from the lecture that these predictors are categorical (in particular binary).

In order to process these values we convert them to the data type `boolean`, i.e., we replace the `String` objects in the columns `default` and `student` by `Boolean`'s.
There are a lot of ways to accomplish this task; the easiest might be

    D.replace(to_replace='No', value=False, inplace=True)
    
**Task**: Replace every 'No' and 'Yes' in the `DataFrame` by the values `False` and `True`, resp.

<!-- grade_id: cell-cc9a0df3e56d8dc7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e19e911aa04de5e9, points: 1, grade: true, solution: false -->
```python
assert D.student.dtype == 'bool'
assert D.default.dtype == 'bool'
```

### Answer the following questions!

Store your answers in the given variables

**Question A**: How many students belong to the data set?

<!-- grade_id: cell-45a0fd6410a6a3f1, points: 0, grade: false, solution: true -->
```python
# answer_A = 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3aa830790df4955d, points: 1, grade: true, solution: false -->
```python
assert 'answer_A' in locals()
```

**Question B**: What is the mean **balance** of all samples?

<!-- grade_id: cell-adbccfcb7c0c4c64, points: 0, grade: false, solution: true -->
```python
# answer_B = 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8f8fa0dd0ba12a6b, points: 1, grade: true, solution: false -->
```python
assert 'answer_B' in locals()
```

**Question C**: What is the mean income of the **students**?

<!-- grade_id: cell-304f5e89cf258b70, points: 0, grade: false, solution: true -->
```python
# answer_C =
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-88cc5b64c139ad51, points: 1, grade: true, solution: false -->
```python
assert 'answer_C' in locals()
```

**Question D**: How many **students** obtain an **income** of more than 20,000.

<!-- grade_id: cell-055cd883f43fd6a8, points: 0, grade: false, solution: true -->
```python
# answer_D = 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ad443f5f82e71c2d, points: 1, grade: true, solution: false -->
```python
assert 'answer_D' in locals()
```

**Question E**: What is the 25% quantile of the predictor **balance**?

<!-- grade_id: cell-d71bb4c3124f7e98, points: 0, grade: false, solution: true -->
```python
# answer_E = 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-65d1394f282d0881, points: 1, grade: true, solution: false -->
```python
assert 'answer_E' in locals()
```

### Plotting the data set

Next, we want to plot both, the `income` and `balance` predictors as boxplots as a function of the `default` status.

**Task**: Complete the plotting command in the following cell. What do you observe?

<!-- grade_id: cell-368b0f91e82b2f93, points: 1, grade: true, solution: true -->
```python
fig, ax = plt.subplots(1,2)
D.boxplot(column='balance',by='default', ax=ax[0]);
# YOUR CODE HERE
raise NotImplementedError()
```

You should observe that it seems that the credit card balance has a large effect on the default status, while the income seems not to predict the default status very well.

Finally, the following cell let's you plot the `default`'s vs. the non-`default`'s of the data set. No task here!

```python
D.plot(y='income', x='balance', kind='scatter',c = D.default, cmap = 'coolwarm', marker='.');
```

### Fitting a logistic regression model
Next, we want to fit a logistic regression model to our data.
Use the `LogisticRegression` function in the module `sklearn.linear_model`.
The behaviour is similar to a `LinearRegression` fit.

You can find the documentation of this function [here](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LogisticRegression.html).
There are a lot of optional arguments, the most important might be the unimpressive looking parameter `C`, which determines the strength of regularization used in the algorithm that solves the maximum likelihood problem.

We will discuss regularization later in the lecture as well as in the labs. For now, it suffices if you keep the following in mind:

**The larger you choose `C`, the less the problem will be regularized.**

**Task**: Fit a logistic regression model that predicts the probability of `default` using `balance` as predictor. You should obtain the following values: $\beta_0: -10.6513$, $\beta_\text{balance}: 0.0055$.

Choose the following optional parameters:
* set the regularization parameter `C = 1e10` (which is the scientific notation of $C = 10^{10}$, and thus very large)
* set the error tolerance to `tol=1e-10`
* set the solver to `solver = 'liblinear'`

in this and the upcoming problems.

<!-- grade_id: cell-22819ddc03810ae9, points: 2, grade: true, solution: true -->
```python
from sklearn.linear_model import LogisticRegression
# YOUR CODE HERE
raise NotImplementedError()
```

**Task**: Store the intercept of the model in a variable `intercept0` and the regression coefficient in the variable `reg_coef0`.
These quantities represent the coeffcients for a linear regression model predicting the log-odds, i.e.

$$
\log \left( \frac{p(x)}{1-p(x)} \right) = \beta_0 + \beta_1 \, x
$$

<!-- grade_id: cell-f5497700d04802be, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6628981a047b598c, points: 1, grade: true, solution: false -->
```python
assert 'intercept0' in locals()
assert 'reg_coef0' in locals()
```

**Task**:
Predict the probability of `default` for a `balance` value of $\$$1,000 and $\$$2,000 and store your answers is the variables `pod_1000` and `pod_2000`, resp.

Use the method `predict_proba` of a `LogisticRegression` model.

**Note**: The model assumes that your data has the same format as your original training data. Therefore, you might have to reshape the input into the correct format.

<!-- grade_id: cell-f089240b8df50cd4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2c1bb68da7976c58, points: 1, grade: true, solution: false -->
```python
assert 'pod_1000' in locals()
assert 'pod_2000' in locals()
```

You should observe, that the probality of default of an individual with a credit card balance of $\$$1,000 is approximately 0.57\%.
The probality of default of an individual with a credit card balance of $\$$2,000 is approximately 58.6\%.

Now, we want to incorporate the predictors `income` and `student` status as well. This can be done easily using the same methods.

Execute the following code cell to train a new logistic regression model.

```python
lr2= LogisticRegression(solver='liblinear', tol=1e-10, C=1e10)
X = D.loc[:,['balance','income','student']]
y = D.loc[:,'default']
reg2 = lr2.fit(X,y)
```

**Task**: Store the intercept of the new model in the variable `intercept_full` as well as the coefficients in variables `beta_balance`, `beta_income`, `beta_student`, resp.

<!-- grade_id: cell-2e01d7749939ac81, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-97b6efdfefdf0bd8, points: 1, grade: true, solution: false -->
```python
assert 'intercept_full' in locals()
assert 'beta_balance' in locals()
assert 'beta_income' in locals()
assert 'beta_student' in locals()
```

**Task**:
What is the default probability of a student and a non-student with a credit card balance of $\$$1,500, an income of $\$$40,000?
Store your answers in the variables `pod_student` and `pod_nonStudent`, resp.

<!-- grade_id: cell-1916a01129726b28, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0b93cc82b179752b, points: 1, grade: true, solution: false -->
```python
assert 'pod_student' in locals()
assert 'pod_nonStudent' in locals()
```

You should observe that a student with a credit card balance of $\$$1,500 and an income of $\$$40,000 has an estimated probability of default of 5.8\%, while an non-student with the same balance and income has a probability of default of 10.5\%.

```python

```