Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below.

Rename this problem sheet as follows:

    ps{number of lab}_{your user name}_problem{number of problem sheet in this lab}
    
for example
    
    ps4_blja_problem1

Submit your homework within one week until next Monday, 9 a.m.

```python
NAME = ""
EMAIL = ""
USERNAME = ""
```

---

# Introduction to Data Science
## Lab 4: Further aspects of linear regression

### Part A - Limitations of the t-test

In this notebook, we investigate the limitations of a single-variable **t-test** for the predictor coefficients $\beta$ in a linear regression setting.
Recall the following statements from the lecture (Slide 38 in 03_IntroDS_Regression.pdf):
* Does a single small $p$-value indicate at least one variable relevant? No.
* Example: $p=100$, $H_0 : \beta_1 = \dots = \beta_p = 0$ true. Then by chance, $5\%$ of $p$-values below $0.05$. Almost guaranteed that $p<0.05$ for at least one variable by chance.
* Thus, for large $p$, looking only at $p$-values of individual $t$-statistics tends to discover spurious relationships.

In what follows, we use slightly different values than in the above mentioned example, setting $n = 100$ and $p = 20$.

```python
import numpy as np

# Set parameters n (number of training samples) and p (number of predictor variables)
n = 100
p = 20
```

For this purpose, we generate random uncorrelated input and output vectors.

**Task**: Write the function `drawSample` that generates **uniformly distributed** arrays of random variables
* $X$ should be of size (n, p+1) with values in $[0,1]$; the first column is reserved for the intercept and should contain a only ones
* $y$ should be of size (n,) with values in $[-0.5,0.5]$

<!-- grade_id: cell-8b1fa66b64d265fb, points: 0, grade: false, solution: true -->
```python
def drawSample(n,p):
    """ This function draws a
    sample for our experiment. """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return (X,y)
```

<!-- grade_id: cell-c21e8f578a5dd23f, points: 1, grade: true, solution: false -->
```python
assert drawSample(40,4)[0].shape == (40,5), 'Wrong shape of X'
assert drawSample(40,4)[1].shape == (40,), 'Wrong shape of y'
assert all(drawSample(40,4)[0][:,0]==1), 'Check the first column of X'
assert drawSample(40,4)[1].min() > -0.5 and drawSample(40,4)[1].max() < 0.5, 'Wrong range of y'
assert drawSample(40,4)[0].min() > 0 and drawSample(40,4)[0].max() <= 1, 'Wrong range of X'
```

The following function computes single-variable t-statistics for the model
$$ y \approx X \beta $$
whose parameters $\beta \in \mathbb{R}^{p+1}$ are estimated via
$$ \hat \beta = (X^\top X)^{-1} X^\top y. $$

<!-- grade_id: cell-888372a43aa00783, points: 0, grade: false, solution: true -->
```python
from scipy.stats import t

def printTStatistic(X, y, p_threshold = 0.10, print_table=True):
    n, m = X.shape
    p = m - 1

    # Invert X^T * X
    V = np.linalg.inv((X.T).dot(X))
    

    # Compute regression coefficients beta
    beta = V.dot( X.T.dot(y) )

    # Extract diagonal of matrix (X^T * X)^-1
    v = V.diagonal()

    # Predict y using beta
    y_pred = X.dot(beta)

    # Compute estimate of sigma
    sigma_hat = np.sqrt( 1./(n-p-1) * np.power(y - y_pred,2).sum() )

    # Compute the standard errors
    SE = np.sqrt(v) * sigma_hat

    # Compute the values of the t-statistic
    t_vals = beta / SE

    # Compute the corresponding p values
    p_vals = 2*t.cdf(-np.absolute(t_vals), n-p-1)

    if print_table:
        
        # Print header
        print('|  Coefficient  | Estimate |    SE    | t-statistic |  p-value  | p < %4.2f |' % p_threshold)
        print('----------------------------------------------------------------------------')
        
        # Print 
        for i in range(p+1):
            pval = p_vals[i]
            if pval < 0.0001:
                pval_str = '< 0.0001'
            else:
                pval_str = '  %5.4f' % pval
            print('|    beta_%02d    |  %6.3f  |  %6.4f  |    %5.2f    | %s  |     %d    |' % (i, beta[i], SE[i], t_vals[i], pval_str, pval < p_threshold))
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

**Task**: Test the function `printTStatistic` using an example drawn with your function `drawSample`.

<!-- grade_id: cell-92e5e11d02c2543d, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now, we want to find out, how many predictor variables are statistically significant for a threshold of $0.10$ in our setting with `n = 100` and `p = 20`.

**Task**: Expand the function `printTStatistic` from above. It should **return the proportion of significant predictor variables** at a certain threshold `p_threshold`. Test it using the example below; execute the next cell multiple times (by hitting `Ctrl + Enter`).

<!-- grade_id: cell-cec12dcb08dc21c7, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Task**: Write a small script that carries out the experiment `1000` times and computes the mean proportion of significant values in our experiment. It should be around `p_threshold`.

**Hint 1**: Use the keyword argument `print_table` to suppress the printing of the tables.

**Hint 2**: You can collect the returned values in a list initialized by `vals = []`. You can append a new value `new_val` using `vals.append(new_val)`.

<!-- grade_id: cell-900fd83ad0f07f69, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Part B: "Nonlinear" linear regression

The goal of this problem is to approximate given data points $(x_i,y_i)$ for $i=1,\ldots,n$ by polynomials of degree $p$.
This can be done by solving the linear regression problem:

$$
 y_i \approx \beta_0 + \beta_1 \, x_i + \beta_2 \, x_i^2 + \ldots + \beta_p \, x_i^p
$$

By splitting our data into a training and test data set, we want to illustrate graphically the problem of overfitting.

**Task**: Define the 'unknown' function

$$
f(x) = \sin(10 \, x) + 5 \, \cos(3 \, x)
$$

using `numpy`.

<!-- grade_id: cell-f885a94e9c772a30, points: 0, grade: false, solution: true -->
```python
import numpy as np

# Define the 'unknown' function f
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-97535adf15a536a9, points: 1, grade: true, solution: false -->
```python
assert(np.abs(f(np.pi)+5) < 1e-8)
assert(np.abs(f(np.pi/2)) < 1e-8)
assert(np.abs(f(0)-5) < 1e-8)
assert(np.abs(f(2) - 5.71) < 1e-2)
```

**Task**: Generate a uniformly distributed random vector `x` of size `n = 200`.

<!-- grade_id: cell-a3beb882c6199b0e, points: 0, grade: false, solution: true -->
```python
# Set random seed to make random variables 'predictable'
np.random.seed(0)

# Generate uniformly distributed data samples over [0,1)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1c0f265beb29c4ea, points: 1, grade: true, solution: false -->
```python
assert n == 200
assert np.abs(x.mean() - 0.5004377979051402) < 1e-8
assert x.shape == (200,)
```

**Task**: Determine the vector `y` in the following way

$$
y_i = f(x_i) + \varepsilon \, \eta_i
$$

with $\eta_i$ standard-normal distributed and $\varepsilon = 1$.

<!-- grade_id: cell-a98eacb6c66093ab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f632b8c0c5988c29, points: 1, grade: true, solution: false -->
```python
assert y.shape == (200,)
assert np.abs(y.mean() - 0.2748887916140714) < 1e-8
```

**Task**: Generate one figure with the following data:
* mark the **data points** $(x_i,y_i)$ as black circles
* draw the **population line** (the line representing the *unknown* function $f$) as a red solid line
* draw the **regression line** for a fitted polynomial with polynomial degree `p = 20` as a blue dashed line

**Hint**: Use the functions `np.polyfit` and `np.polyval` to determine the regression line.

<!-- grade_id: cell-374277908e484443, points: 4, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt
plt.rcParams['figure.figsize'] = (15,8)

fig = plt.figure(1, clear = True)

# YOUR CODE HERE
raise NotImplementedError()

```

Split the dataset $(x,y)$ into a training and test set using `np.split`
- the training set should contain `ntrain` samples
- the test set should contain `n - ntrain` samples

Choose `ntrain = 80`.

<!-- grade_id: cell-8f48e7bce19fd7cb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7e8109022605f809, points: 1, grade: true, solution: false -->
```python
assert(ntrain == 80)
assert(xtrain.shape == (80,))
assert(xtest.shape == (120,))
assert(ytrain.shape == (80,))
assert(ytest.shape == (120,))
assert np.hstack((xtrain,xtest)).mean() == x.mean()
```

Now we want to fit polynomial models with varying polynomial degrees ($p= 0,\ldots,20$).
As a quality measure, we store the training MSE (mean squared error) and the test MSE.

**Note**: You can ignore the `RankWarning`s!

<!-- grade_id: cell-fc2d52ce5ecc280d, points: 0, grade: false, solution: true -->
```python
def computeMSE(y, fhatx):
    " This function returns the mean squared error between x and y."
    return np.mean(np.power(y-fhatx,2))

# Initialize lists that contain test and training mean squared errors
MSEtrain = []
MSEtest = []

# Set range for different degrees
# YOUR CODE HERE
raise NotImplementedError()

for j in deg_range:
    
    # Fit polynomial of degree 'j' on training data
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Append test and training mse to according list
    # YOUR CODE HERE
    raise NotImplementedError()


```

**Task**: Generate one figure that contains
- the test mse in a logarithmic plot as a blue dashed line
- the training mse in a logarthmic plot as a red solid line

against the polynomial degree.
You should use the function `plt.semilogy` and set meaningful `label`s.

<!-- grade_id: cell-ca2617631310e831, points: 4, grade: true, solution: true -->
```python
fig = plt.figure(2, clear=True)
# YOUR CODE HERE
raise NotImplementedError()
plt.legend()
plt.xlabel("Polynomial degree")
plt.ylabel("MSE")
plt.show()
```